```python
# Initialize Otter
import otter
grader = otter.Notebook("JHW3.ipynb")
```

<img src="images/flips.jpg" style="float: right; width: 30%">

<h1> Jupyter Homework Week 3 --- Math 32, Summer 2024 </h1>


<h2> The Binomial and Geometric Disributions </h2>


In this notebook, we'll use some readily-available packages that will help us calculate probabilities of the [binomial distribution](https://en.wikipedia.org/wiki/Binomial_distribution), and also sample from the [geometric distribution](https://en.wikipedia.org/wiki/Geometric_distribution). We'll use our sampling of the geometric distribution to approximate an answer to a question in your textbook (exercise **4.12** in *Dekking, et al*). 

```python
# import the requisite packages
from scipy.stats import binom
from scipy.stats import geom
import numpy as np
import scipy
```

<h3> Binomial distributions </h3>

<h4> Part 1: Getting a sense of the binomial distribution </h4>

The binomial distribution has probability mass function


$$p(k) = \binom{n}{k} p^k (1 - p)^{n - k} \ \ \ \ \ $$ for $k=0,1,...,n$

We have seen in class how this distribution models the *number of successes* in a sequence of $n$ independent and identically distributed trials, with probability of success $p$. We can get the values of it's pmf and cdf using [scipy.stats.binom](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.binom.html), an object from scipy's package of statistical functions.

```python
# example usage of binom
k = 1
n = 10
p = 0.2
print("n =",n,"p =",p)
print("Probability of exactly one success:", binom.pmf(k,n,p))
```

```python
print("Probability of at least one success: ",binom.pmf(0,n,p)+binom.pmf(1,n,p))
print("Using the cdf function also works:   ",binom.cdf(1,n,p))
```

## Question 1: 

- Suppose you flip a fair ($p = 0.5$) coin $100$ times. How likely is it that you'll see exactly $50$ heads? What is the probability that you'll see between $45$ and $55$ heads (including $45$ and $55$)?

Hint: Use `binom` function as in the above example

<!-- BEGIN QUESTION -->



```python
# Fill in the ... with your answer

n= ...
p = ...
probability_exactly_50 = ...


probability_between_45_and_55 = ...

print(probability_exactly_50)
print(probability_between_45_and_55)
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<h3> Geometric distribution </h3>

<h4> Part 2: Getting a sense of the geometric distribution </h4>

The geometric distribution has the probability mass function


$$P(X = k) = (1 - p)^{k - 1}p \ \ \ \ \  \ \ \ \ \ \ \ \ \ \ k=1,2,...$$

We have seen in class how this distribution models the *time of the first success* in a sequence of independent and identically distributed trials. For example, suppose we flip a coin until we get our first *Heads*. If the sequence is

$$TTTH,$$

then $X=4$ since it took *four total flips* to get our first *Heads*. If the experiment results instead in 

$$TTTTH,$$

then $X = 5$.


We can **sample** from the distribution using numpy's `np.random.geometric` function, which you can read about [in the documentation](https://numpy.org/doc/stable/reference/random/generated/numpy.random.geometric.html).

### Question 2:

In this part of the weekly notebook, you'll use simulation to answer a problem from the textbook (Exercise 4.12):

> You and a friend want to go to a concert, but unfortunately only one ticket is still available. The man who sells the tickets decides to toss a coin until heads appears. In each toss heads appears with probability $p$, where $0 < p < 1$, independent of each of the previous tosses. If the number of tosses needed is odd, your friend is allowed to buy the ticket; otherwise you can buy it. Would you agree to this arrangement?

Your goal is to 

- Sample a large number of times from the geometric distribution using `np.random.geometric`.
- Record the number of times the count is *even* using `even_count`.
- Store the frequency of *even* in the variable `probability_even_tosses`, this should estimate the probability of observing an even number of tosses.

The function should take two parameters: the probability of success `p` and the number of simulations `num_simulations`.

```python
def probability_even_tosses(p, num_simulations=10000):
    """
    Calculate the probability of observing an even number of tosses
    using the np.random.geometric function.
    Parameters:
    - p: Probability of success (getting heads) in each toss.
    - num_simulations: Number of simulations to run (default is 10,000).
    Returns:
    - Probability of observing an even number of tosses.
    """
    
    even_count = ...
    for _ in range(num_simulations):
        # Generate the number of tosses to get the first success
        num_tosses = ...
        # Check if the number of tosses is even
        ...
            ...
    # Calculate the frequency of observing an even number of tosses
    frequency = ...
    return frequency
  
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 5)
```
>>> assert round(probability_exactly_50, 4) == 0.0796
>>> assert round(probability_between_45_and_55, 4) == 0.7287
```


Tests for q2 (points: 10)
```
>>> def test_probability_even_tosses(p):
...     result = probability_even_tosses(p, num_simulations=10000)
...     assert isinstance(result, float), 'The result should be a float.'
...     assert 0 <= result <= 1, 'The result should be a probability value between 0 and 1.'
...     result = probability_even_tosses(p, num_simulations=100000)
...     expected_value = p * (1 - p) / (1 - (1 - p) ** 2)
...     tolerance = 0.005
...     assert abs(result - expected_value) <= tolerance, f'The result should be close to {expected_value}.'
>>> 
>>> def test_probability_even_tosses_hidden(p):
...     result = probability_even_tosses(p, num_simulations=100000)
...     assert isinstance(result, float), 'The result should be a float.'
...     assert 0 <= result <= 1, 'The result should be a probability value between 0 and 1.'
...     result = probability_even_tosses(p, num_simulations=100000)
...     expected_value = p * (1 - p) / (1 - (1 - p) ** 2)
...     tolerance = 0.005
...     assert abs(result - expected_value) <= tolerance, f'The result should be close to {expected_value}.'
```
