```python
# Initialize Otter
import otter
grader = otter.Notebook("JHW5.ipynb")
```

<h2> Jupyter Homework Week 5 --- Long-term averages </h2>

<img src="images/lln.png" style="float: right; width: 30%">

In this week's notebook, you'll explore data and understand how the averages tend towards their limits. We'll generate data from different distributions and see how it settles down towards a particular value over time, by repeatedly computing averages.

```python
# import the requisite packages
import numpy as np
from math import factorial
```

<h3> Process </h3>



<img src="images/uniform.png" style="float: right; width: 30%">

The uniform distribution $\operatorname{Unif}(0, 1)$ has mean $1/2 = 0.5$ and variance $1/12 = 0.0833...$. We can sample from this distribution using numpy's `np.random.uniform()` method.

If $X_1,X_2,...,X_n$ are independent samples of a uniform distribution between $a=0$ and $b=1$, the sample mean is defined as

$$\overline{X}_{n} = \frac{X_1 + X_2 + \cdots + X_{n}}{n}.$$


The sample mean $\overline{X}_{100}$ of $100$ samples will tend to be *closer* to the true mean $\mu = 0.5$ then any one sample. In fact, the more and more samples we take, the sample mean will converge to the true mean. We will explore how the variance of $\overline{X_n}$ goes to $0$ as $n$ gets larger (we will use test $n=10$, $n=100$, and $n=1000$.)

<h3> Question 1 </h3>


* Create a function that returns the sample mean of $n$ samples, where $n$ is an input to the function. You can use `samples = np.random.uniform(size = n)` to create the samples, and `samples.mean()` to return the mean.
  


```python
def sample_mean(n):
    # sample n times from a uniform distribution
    samples = ...
    # return the mean of the samples
    ...
```

```python
print("The sample mean of n=10 samples:",sample_mean(10))
print("The sample mean of n=100 samples:",sample_mean(100))
print("The sample mean of n=1000 samples:",sample_mean(1000))
```

```python
grader.check("q1")
```

### Question 2

* If $X$ has the original $\operatorname{Unif}(0, 1)$ distribution, then there is an $80\%$ chance that $X$ differs from its mean by at least $0.1$: $$P(|X - 0.5| \ge 0.1) = 80\%.$$
Estimate this quantity for $\overline{X}_{10}$,$\overline{X}_{100}$, and $\overline{X}_{1000}$ replacing $X$. Use a for-loop to sample $\overline{X}$ many times. Use a variable called `count` to keep track of how often $|X-0.5|\geq 0.1$ happens.

<img src="images/xbars.png" style="float: center; width: 50%">

```python
def calculate_frequency(n):
    # use a large number of samples (e.g. 10K)
    num_samples = ...
    
    # intialize a count variable
    count = ...
    
    
    for _ in range(num_samples):
        # simulate using 'sample_mean'
        xbar = ...
        # add 1 to count if xbar <= 0.4 or xbar>=0.6
        ...
            ...

    # calculate the frequency
    frequency = ...
    return frequency


print("n=10: An estimate for the probability:",calculate_frequency(10))
print("n=100: An estimate for the probability:",calculate_frequency(100))
print("n=1000: An estimate for the probability:",calculate_frequency(1000))
```

```python
grader.check("q2")
```

What do you notice about the probabilities? Do they increase/decrease as $n$ gets larger? Does this match your intuition?

```python
# Write your answer here


```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: [2, 2, 2])
```
>>> def irwin_hall_check(n, a, num_samples=10000):
...     samples_n = np.array([sample_mean(n) for _ in range(num_samples)])
...     prob_error = (1 - 1 / factorial(n + 1) * (a * n) ** (n + 1)) ** num_samples
...     return samples_n.min() > a
>>> n, a = (10, 0.1)
>>> assert irwin_hall_check(n, a)
```
```
>>> def irwin_hall_check(n, a, num_samples=10000):
...     samples_n = np.array([sample_mean(n) for _ in range(num_samples)])
...     prob_error = (1 - 1 / factorial(n + 1) * (a * n) ** (n + 1)) ** num_samples
...     return samples_n.min() > a
>>> n, a = (100, 0.01)
>>> assert irwin_hall_check(n, a)
```
```
>>> def irwin_hall_check(n, a, num_samples=10000):
...     samples_n = np.array([sample_mean(n) for _ in range(num_samples)])
...     prob_error = (1 - 1 / factorial(n + 1) * (a * n) ** (n + 1)) ** num_samples
...     return samples_n.min() > a
>>> n, a = (1000, 0.001)
>>> assert irwin_hall_check(n, a)
```


Tests for q2 (points: [2, 2])
```
>>> assert np.isclose(calculate_frequency(10), 0.28074, atol=0.2)
```
```
>>> assert np.isclose(calculate_frequency(100), 0.000511, atol=0.01)
```
